## Vectorizing code ##

It is important to make sure that the code is efficient. This means that operations should be vectorized where possible, instead of relying on `for` loops.

This is particularly important when doing the exercises, as they are automatically graded. If the code is not efficient, some of the tests that are used to evaluate your solution might exceed time limits, such that you will not get full points.

Here is an example: Given a matrix $K \in \mathbb R^{n \times n}$ and a vector $\mathbf a \in \mathbb R^n$ you should write code to calculate
$$
\sum_{i=1}^n \sum_{j=1}^n a_i a_j K_{ij}
$$

```python
def matrix_expectation_with_loops(a, K):
    """
    Return the sum of a[i] * a[j] * K[i,j] for all i, j using two for loops.
    """
    result = 0
    for i in range(len(a)):
        for j in range(len(a)):
            result += a[i] * a[j] * K[i,j]
    return result
```

```python
def matrix_expectation_vectorized(a, K):
    """
    Return the sum of a[i] * a[j] * K[i,j] for all i, j using numpy's matrix multiplication.
    """
    return a.T @ K @ a
```

```python
import time
import numpy as np
np.random.seed(42)

n = 2000
a = np.random.randint(100, size=n)
K = np.random.randint(100, size=(n,n))

t0 = time.time()
print(matrix_expectation_with_loops(a, K))
t1 = time.time()
print('Time for loops = {:.3f} seconds.'.format(t1-t0))  # This is too slow.

t0 = time.time()
print(matrix_expectation_vectorized(a, K))
t1 = time.time()
print('Time for vectorized = {:.3f} seconds.'.format(t1-t0))  # This is fast.
```

<!-- grade_id: cell-9fcaf33641894961, points: 0, grade: false, solution: true -->
```python
def my_solution(a, K):
    """Return the sum of a[i] * a[j] * K[i,j] for all i, j using vectorization.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: tests_solution, points: 1, grade: true, solution: false -->
```python
import time
import numpy as np
np.random.seed(42)

n = 100
a = np.random.randint(100, size=n)
K = np.random.randint(100, size=(n,n))

solution = 1261634040
time_tolerance = 10  # in seconds

t0 = time.time()
np.testing.assert_equal(my_solution(a, K), solution)
t1 = time.time()
total = t1 - t0
assert total < time_tolerance, 'Time limit exceeded: %f' % total

n = 5000
a = np.random.randint(100, size=n)
K = np.random.randint(100, size=(n,n))

solution = 3082209422285
time_tolerance = 10  # in seconds

t0 = time.time()
np.testing.assert_equal(my_solution(a, K), solution)
t1 = time.time()
total = t1 - t0
assert total < time_tolerance, 'Time limit exceeded: %f' % total
```